# Question 1

In this question, you'll write some coding that performs string manipulation. This is pretty much your warm-up.

The first cell block is where you put your answer. There may already be some code there, in which case **don't remove it**; just add what is asked for.

All the cells under the first one are autograder cells; don't edit those! But you can run them (using the Run button at the top) to test your own code! If they don't spit out any error messages, then your code works!

### Part A

What's your favorite *positive* number? Create a variable named `favorite_number` and assign it to be your favorite *positive* number.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1a_test, points: 6, grade: true, solution: false -->
```python
print("My favorite number is: " + str(favorite_number))
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-933f5d1987d3> in <module>()
----> 1 print("My favorite number is: " + str(favorite_number))

NameError: name 'favorite_number' is not defined
```

### Part B

Print out a famous quote! In the code below, create two variables:
 - one named `famous_person`,
 - the other named `their_quote`,
 
and assign them to be the name of some famous person and a quote that they said, respectively. Remember: these are *strings*!

<!-- grade_id: q1b, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1b_test, points: 6, grade: true, solution: false -->
```python
print("{}, at age {}, said:\n\n\"{}\"".format(famous_person, favorite_number, their_quote))

assert len(famous_person) > 0
assert len(their_quote) > 0
```

### Part C

You're working late on a homework assignment and have copied a few lines from a Wikipedia article. In your tired stupor, your copy/paste skills leave something to be desired. Rather than try to force your mouse hand to stop shaking, you figure it's easier to write a small Python program to strip out errant whitespace from your copying-pasting.

**Leaving the code below as-is**, **add more code** that strips away whitespace in the variables `line1`, `line2`, and `line3`.

In other words, reassign each of the three variables (keeping their names the same, and their respective content otherwise unaltered) so that the leading and trailing whitespace on both ends is removed.

```python
line1 = 'Python supports multiple programming paradigms, including object-oriented, imperative\n'
line2 = '  and functional programming or procedural styles. It features a dynamic type\n'
line3 = ' system and automatic memory management and has a large and comprehensive standard library.\n '
```

<!-- grade_id: q1c, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1c_test, points: 6, grade: true, solution: false -->
```python
assert line1[-1] != '\n'
assert line2[0] != line2[1] != ' '
assert line2[-1] != '\n'
assert line3[0] != ' '
assert line3[-1] != '\n'
```

### Part D

You discover that there are numbers in the text you'd like to be able to parse out for some math down the road. After stripping out the trailing whitespace, convert the numbers to their proper numerical form.

Assign them to the variables `num1`, `num2`, and `num3` respectively.

```python
line1n = ' 495.59863 \n'
line2n = '\t134  '
line3n = '\n\t -5.4 \t'
```

<!-- grade_id: q1d, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1d_test, points: 6, grade: true, solution: false -->
```python
assert num1 > 495 and num1 < 496
assert num2 == 134
assert num3 > -6 and num3 < -5
```

### Part E

Take the `number` below, find its square root, convert it to a string, and then assign it to a variable named `str_version`.

You **must use the correct arithmetic operator for the square root, as well as the correct casting function for the string conversion**.

**No imports allowed!**

<!-- grade_id: q1e, points: 0, grade: false, solution: true -->
```python
number = 3.14159265359
```

```python

```

<!-- grade_id: q1e_test, points: 6, grade: true, solution: false -->
```python
assert str_version == "1.7724538509055743"
```

# Question 2

In this question, we'll dive a bit deeper into some arithmetic operations.

### Part A

Five numbers are given below. Keeping them in that same order, compute a single result with Python addition, multiplication, subtraction, and division, respectively.

Put another way--add the first and second numbers, multiply the second and third numbers, subtract the third and fourth numbers, and divide the fourth and fifth numbers. Don't use any parentheses; let order of operations take care of things. Do this in a single statement, and store the result in the `result` variable below.

```python
n1 = 0.4589
n2 = 13
n3 = -59
n4 = 8.43333333
n5 = 2.1

result = -1
```

```python

```

```python
import numpy as np
np.testing.assert_allclose(result, -770.5569730142857)
```

### Part B

Using the same five numbers as before, this time use parentheses around your arithmetic operation to enforce that the *addition* and *subtraction* operations should happen before the others. Again, store your result in the variable `result`.

```python
n1 = 0.4589
n2 = 13
n3 = -59
n4 = 8.43333333
n5 = 2.1

result = -1

```

```python

```

```python
import numpy as np
np.testing.assert_allclose(result, -432.1802333119699)
```

### Part C

Repeat the steps you took in Part B, this time casting each variable to an integer data type before performing any computation. Store your answer in the variable `result`.

```python
n1 = 0.4589
n2 = 13
n3 = -59
n4 = 8.43333333
n5 = 2.1

result = -1

```

```python

```

```python
assert result == -435.5
```

### Part D

Two numbers are given below. When you perform the operation `n1 / n2`, what is the quotient? Store this quantity in the variable named `quotient`.

```python
n1 = 20344983945038.48374839
n2 = 984374.783894273

quotient = -1

```

```python

```

```python
import numpy as np
np.testing.assert_allclose(quotient, 20667924.735487375)
```

### Part E

You've seen in lecture how, when two numbers in Python are divided, their result is a float. What if you *really* wanted the output of division to be an integer? See if you can make that happen by modifying the following single line of code.

(You are **not** allowed to add any numbers or additional lines of code!)

```python
integer_quotient = 32 / 16
```

```python

```

```python
assert type(integer_quotient) == int
```

# Question 3

In this question, we'll focus on looping.

### Part A

In the question below, you'll write code that generates a new list, `reverse`, that is the same as the input list except with its elements backwards relative to the original.

As an example, if the input list is `[1, 2, 3]`, the output list `reverse` should be `[3, 2, 1]`.

**You have to use loops** to answer this question! `[::-1]` is a neat trick that you should definitely use, and `reversed` is a great function you should also definitely use, but you just can't use them here.

```python
def reverse_list(inlist):
    reverse = []
    
    ### BEGIN SOLUTION
    
    iterator = len(inlist) - 1 
    
    while iterator >= 0:
        
    
    ### END SOLUTION    
    
    return reverse
```

```python
t1 = [1, 2, 3]
t2 = [3, 2, 1]
assert t2 == reverse_list(t1)
```

```python
t1 = [2, 6, 52, 64, 3.4, "something", 3, 6.2]
t2 = [6.2, 3, "something", 3.4, 64, 52, 6, 2]
assert t1 == reverse_list(t2)
```

### Part B

This time, you'll implement a method for computing the average value of a list of numbers. Remember how to compute an average: add up all the values, then divide that sum by the number of values you initially added together.

Example Input: [2, 1, 4, 3]

Example Output: 2.5

Aside from `len()`, you **cannot** use built-in functions like `sum()` or import any packages to help you out. You have to write this on your own!

```python
def average(numbers):
    avg_val = 0.0
    
    ### BEGIN SOLUTION
    
    ### END SOLUTION
    
    return avg_val
```

```python
import numpy as np

inlist = np.random.randint(10, 100, 10).tolist()
np.testing.assert_allclose(average(inlist), np.mean(inlist))
```

```python
inlist = np.random.randint(10, 1000, 10).tolist()
np.testing.assert_allclose(average(inlist), np.mean(inlist))
```

### Part C

In this question, the function below `first_negative` takes a list of numbers (creatively-named `numbers`). Your goal is to write a `while` loop that searches for the **first negative number** in the list.

You should set the variable `num` equal to that negative number.

In this question the `numbers` list will always contain at least one negative number.

```python
def first_negative(numbers):
    num= numbers[0]
    
    ### BEGIN SOLUTION
            
    ### END SOLUTION

    return num
```

```python
inlist1 = [1, 4, -22, -39, 1, 8, 27, 6, 27, -12]
a1 = -22
assert a1 == first_negative(inlist1)
```

```python
inlist2 = [-29, -1, 49, -33, 49, -40, -7, 33, 36, 41]
a2 = -29
assert a2 == first_negative(inlist2)
```

```python

```