# Question 1

In this question, we're looking at some of the basic operations of native Python data structures and iterations.

### Part A

In this part, you'll be given three lists: X coordinates, Y coordinates, and Z coordinates. Therefore, the first data point in this 3D space is `(X[0], Y[0], Z[0])`; the second one is `(X[1], Y[1], Z[1])`, and so on. You will need to use these coordinates and compute the *magnitudes* of the vectors they form.

Recall that a 2D, 3D or any-D point in space can also be represented as a vector, and vectors have *lengths*, or magnitudes. This length can be computed using the Euclidean distance function:

$$
d = \sqrt{x^2 + y^2 + z^2}
$$

Recall that you can compute exponents in Python using the `**` operator.

Append these magnitudes to the list `magnitudes`.

**No imports allowed!**

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
def compute_3dmagnitudes(X, Y, Z):
    magnitudes = []
    
    ### BEGIN SOLUTION
    
    
            
    ### END SOLUTION
    
    return magnitudes
```

<!-- grade_id: q1a_test1, points: 10, grade: true, solution: false -->
```python
import numpy.testing as t

x1 = [1, 2, 3]
y1 = [4, 5, 6]
z1 = [7, 8, 9]
a1 = [8.12403840463596, 9.643650760992955, 11.224972160321824]
t.assert_allclose(compute_3dmagnitudes(x1, y1, z1), a1)
```

<!-- grade_id: q1a_test2, points: 10, grade: true, solution: false -->
```python
x2 = [ 30.48325629,  43.69168085,   1.77825571,  79.51618164, 18.97152147,  71.22477376,  98.29697101,  41.97578185, 10.61862458,  93.50326655]
y2 = [ 98.97761722,  75.44498495,  41.40658662,  76.24025762, 12.64196712,  99.80696952,  27.68903488,   5.02806986, 92.34989725,  95.81159194]
z2 = [ 12.72332579,  22.04915914,  48.91187418,  24.43381236, 63.35006599,  97.43294649,  42.61849556,  39.79723282, 44.2374771 ,  16.67458157]
a2 = [104.34404939309736,89.92816104193963,64.10958621524445,112.83798656043598,67.32732595675537,156.61283032897592,110.65854384603286,58.06089465866948,102.94762304312597,134.91020597468528]
t.assert_allclose(compute_3dmagnitudes(x2, y2, z2), a2)
```

### Part B

In this question, you're given two lists: keys and values. Write a loop that uses the `zip` function in order to create a dictionary which contains the specific keys and values.

For example, the function `lists_to_dict( ["one", "two", "three"], [1, 2, 3] )` should return a dictionary that looks like:

`{"one": 1, "two": 2, "three": 3}`

It doesn't matter if, when printed, the key/values are in a different order, as long as the correct keys are paired with the correct values.

**Sorry, using `dict()` is not allowed.** It's a great function that would answer this question in one line, but you can't use it here.

<!-- grade_id: q1b, points: 0, grade: false, solution: true -->
```python
def lists_to_dict(keys, values):
    new_dict = {}
    
    ### BEGIN SOLUTION
    
    
            
    ### END SOLUTION   
    
    return new_dict
```

<!-- grade_id: q1b_test1, points: 10, grade: true, solution: false -->
```python
k1 = ["1", "2", "3"]
v1 = ["one", "two", "three"]
a1 = set({(k1[0], v1[0]), (k1[1], v1[1]), (k1[2], v1[2])})
assert a1 == set(lists_to_dict(k1, v1).items())
```

<!-- grade_id: q1b_test2, points: 10, grade: true, solution: false -->
```python
k2 = ["wat1", "wat2", "wat3", "wat4", "wat5"]
v2 = [1, 2, 3, 4, 5]
a2 = set({(k2[0], v2[0]), (k2[1], v2[1]), (k2[2], v2[2]), (k2[3], v2[3]), (k2[4], v2[4])})
assert a2 == set(lists_to_dict(k2, v2).items())
```

# Question 2

In this question, we'll look at using conditionals to change the behavior of code.

### Part A

In this question, you'll write a method that takes a list of numbers [0-9] and returns a corresponding list with the "ordinal" versions. That is, if you see a 1 in the list, you'll create a string "1st". If you see a 3, you'll create a string "3rd", and so on.

For example, if you receive `[2, 1, 4, 3, 4]` as input, you should create a list of strings that looks like this: `["2nd", "1st", "4th", "3rd", "4th"]`.

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def return_ordinals(numbers):
    out_list = []
    
    ### BEGIN SOLUTION
    
    
    
    
    
    
    
    
            
    ### END SOLUTION    

    return out_list
```

<!-- grade_id: q2a_test1, points: 5, grade: true, solution: false -->
```python
inlist = [5, 6, 1, 9, 5, 5, 3, 3, 9, 4]
outlist = ["5th", "6th", "1st", "9th", "5th", "5th", "3rd", "3rd", "9th", "4th"]
for y_true, y_pred in zip(outlist, return_ordinals(inlist)):
    assert y_true == y_pred.lower()
```

<!-- grade_id: q2a_test2, points: 5, grade: true, solution: false -->
```python
inlist = [7, 5, 6, 6, 3, 5, 1, 0, 5, 2]
outlist = ["7th", "5th", "6th", "6th", "3rd", "5th", "1st", "0th", "5th", "2nd"]
for y_true, y_pred in zip(outlist, return_ordinals(inlist)):
    assert y_true == y_pred.lower()
```

### Part B

In this question, you'll write code that computes the median of a sorted list of numbers. You can assume the list you receive is already sorted in ascending order (least to greatest).

If your input list is `[1, 1, 2, 4, 7, 7, 8]`, then your output should be `4`. Recall the rule about median: if you get a list with an *even* number of elements, you should return the average of the two middle ones.

Store your answer in the variable `med_num`.

**No imports allowed!**

<!-- grade_id: q2b, points: 0, grade: false, solution: true -->
```python
def median(numbers):
    med_num = 0
    
    ### BEGIN SOLUTION
    
    
    
    
            
    ### END SOLUTION    
    
    return med_num
```

<!-- grade_id: q2b_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(5894835)

l1 = (np.random.random(1000) * 2000) - 1000
l1.sort()
outmed = 1
inlist = l1.tolist()

assert outmed == int(median(inlist))

```

<!-- grade_id: q2b_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(69582)

l2 = (np.random.random(1000) * 2000) - 100
l2.sort()
outmed = 849
inlist = l2.tolist()

assert outmed == int(median(inlist))

### BEGIN HIDDEN TESTS
l22 = (np.random.random(100) * 100) - 50
l22.sort()
om2 = 5
il2 = l22.tolist()
assert om2 == int(median(il2))
### END HIDDEN TESTS
```

### Part C

In this question, you'll write code to find the mode of a list of numbers. Recall that the mode is the number that occurs most frequently. Ties can be broken arbitrarily (meaning you can pick whichever number among those tied for the most frequent).

If your input list is `[5, 1, 3, 1, 2, 5, 1]`, you should return `1`.

Hint: try using a dictionary to maintain the number of times you've seen particular numbers.

Store your answer in the variable `mode_num`.

**No imports allowed!**

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python
def mode(numbers):
    mode_num = 0

    ### BEGIN SOLUTION
    
    
            
        
        
        
        
        
        
        
        
        
        
        
        
    ### END SOLUTION   
    
    return mode_num
```

<!-- grade_id: q2c_test1, points: 5, grade: true, solution: false -->
```python
l1 = [5, 1, 3, 1, 2, 5, 1]
a1 = 1
assert mode(l1) == a1
```

<!-- grade_id: q2c_test2, points: 5, grade: true, solution: false -->
```python
l2 = [1, 2, 3, 1, 2, 3, 1, 2, 3, 4, 4, 4, 4, 4]
a2 = 4
assert mode(l2) == a2

### BEGIN HIDDEN TESTS
l3 = [5, 5, 5, 5, 5, 6, 6, 6, 6, 6, 6, 6, 6]
a3 = 6
assert mode(l3) == a3
### END HIDDEN TESTS
```