# Question 1

Loops, loops, and more loops. And writing basic functions.

### Part A

Write a function:

 - named `add_one`
 - takes 1 argument, a list
 - returns 1 value, a list
 
In this function, you should loop through each element of the argument list, add 1 to it, and then return the new list.

For example, `add_one([1, 2, 3])` should return `[2, 3, 4]`. No imports!

*HINT*: I would highly recommend constructing a *new* list, rather than updating the *old* one.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1a_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(7584)

list1 = np.random.randint(-100, 100, 10)
actual1 = list1 + 1
pred1 = add_one(list1.tolist())
assert set(pred1) == set(actual1.tolist())
```

<!-- grade_id: q1a_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(68527)

list2 = np.random.randint(-100, 100, 100)
actual2 = list2 + 1
pred2 = add_one(list2.tolist())
assert set(pred2) == set(actual2.tolist())
```

### Part B

Write a function that adds an arbitrary integer quantity to every element in a list. This is just like Part A, except instead of adding 1 to each element, you add some positive integer x to each one.

This function should:

 - be named `add_to_list`
 - take 2 arguments: the list, and the number
 - return a list: the new list that contains the added elements
 
In this function, you should loop through each element of the argument list, add the second argument to it, and then return the new list.

For example, `add_to_list([1, 2, 3], 1)` should return `[2, 3, 4]`, just like in Part A. `add_to_list([1, 2, 3], 5)` should return `[6, 7, 8]`. No imports!

*HINT*: I would highly recommend constructing a *new* list, rather than updating the *old* one.

<!-- grade_id: q1b, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1b_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(2846)

list1 = np.random.randint(-100, 100, 10)
num1 = 15
actual1 = list1 + num1
pred1 = add_to_list(list1.tolist(), num1)
assert set(pred1) == set(actual1.tolist())
```

<!-- grade_id: q1b_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(68527)

list2 = np.random.randint(-100, 100, 100)
num2 = np.random.randint(0, 100)
actual2 = list2 + num2
pred2 = add_to_list(list2.tolist(), num2)
assert set(pred2) == set(actual2.tolist())
```

### Part C

Write a function:

 - named `list_of_positive_indices`
 - takes 1 arguments, a list
 - returns 1 value, a list
 
In this function, you should loop through each element of the argument list. If the element is *positive*, then you should append the *index number of the positive element* to a new list. If not, skip that index (i.e. do nothing). Once the loop is finished, return the new list.

For example, `list_of_positive_indices([1, -1, 0, 3])` should return `[0, 3]`. No imports!

*HINT*: Instead of doing a simple `for element in list:` loop, try using the `range()` function in the loop header.

<!-- grade_id: q1c, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(986)

list1 = np.random.randint(-100, 100, 10)
indices1 = np.arange(list1.shape[0])[list1 > 0].tolist()
pred1 = list_of_positive_indices(list1.tolist())
assert set(pred1) == set(indices1)
```

<!-- grade_id: q1c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(578752)

list2 = np.random.randint(-100, 100, 100)
indices2 = np.arange(list2.shape[0])[list2 > 0].tolist()
pred2 = list_of_positive_indices(list2.tolist())
assert set(pred2) == set(indices2)
```

# Question 2

More loops, more functions, and this time with conditionals!

### Part A

Write a function that tests if a number is even.

*HINT*: Recall the modulo operator `%`. This performs division on two numbers, but instead of returning their quotient, it returns the *remainder* of the division. For example, `7 % 3` returns 1, since the quotient is 2 with a remainder of 1.

Your function should:

 - be named `is_even`
 - take 1 argument, a number
 - return `True` if the argument is even, `False` otherwise
 
Make sure you're returning a *boolean*!

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q2a_test1, points: 2.5, grade: true, solution: false -->
```python
assert is_even(3) == False
```

<!-- grade_id: q2a_test2, points: 2.5, grade: true, solution: false -->
```python
assert is_even(123456789) == False
```

<!-- grade_id: q2a_test3, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42) == True
```

<!-- grade_id: q2a_test4, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42.1) == False
```

### Part B

Write a function that finds the *second-smallest* number in a list of numbers. Your function should:

 - be named `second_smallest`
 - take 1 argument, a list of numbers
 - return 1 number, the second-smallest number in the list
 - throw an `Exception` if there are not enough elements in the list
 
Take note of that last point. There's a possibility, after all, that the list your function receives as input will have 1 or even 0 elements in it. How, then, do you return the second-smallest element of a 0-element list? Well, simply, you can't. So instead, you'll raise an error.

Also note: you're the one *raising* the error, so it's *not* your job to catch it! It's the job of the person *using* your function to note the possibility that it will throw an error, and to take appropriate action to handle it gracefully. So if your code includes any `try` or `except` statements, you're doing it wrong!

For example, `second_smallest([1, 2, 3])` should return `2`. Another example `second_smallest([10, 100, 1000, 10000])` should return 100. And `second_smallest([1])` should raise an `Exception`.

You cannot use any built-in sorting functions! Only `range()` and `len()`.

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q2c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(89547)

l1 = np.random.randint(-100, 100, 100)
a1 = np.sort(l1)[1]
np.testing.assert_allclose(a1, second_smallest(l1.tolist()))
```

<!-- grade_id: q2c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(485)

l2 = np.random.randint(0, 1000, 100)
a2 = np.sort(l2)[1]
np.testing.assert_allclose(a2, second_smallest(l2.tolist()))
```

<!-- grade_id: q2c_test3, points: 2.5, grade: true, solution: false -->
```python
try:
    second_smallest([])
except:
    assert True
else:
    assert False
```

<!-- grade_id: q2c_test4, points: 2.5, grade: true, solution: false -->
```python
try:
    second_smallest([100])
except:
    assert True
else:
    assert False
```

# Question 3

This question will focusing on looping and using dictionaries.

### Part A

In this part, you'll write a safe version of testing a dictionary for a specific key and extracting the corresponding element. Normally, if you try to access a key that does not exist, your program crashes with a `KeyError` exception.

Here, you'll write a function which:

 - is named `safe_access`
 - takes 3 arguments: a dictionary, a key, and a *default return value*
 - returns 1 value: the value in the dictionary for the given key, or the default return value
 
The third argument to this function is a *default argument*, whose initial value should be the Python keyword `None`.

Your task in this function is three-fold:

 1. Try to access the value in the dictionary at the given key (wrap this in a `try-except` statement)
 2. If the access is successful, just return the value
 3. If the access is NOT successful (i.e. you caught a `KeyError`), then return the default value instead
 
You must include a `try-except` statement in your function!

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q3a_test1, points: 5, grade: true, solution: false -->
```python
v1 = safe_access({"one": [1, 2, 3], "two": [4, 5, 6], "three": "something"}, "three")
assert v1 == "something"
```

<!-- grade_id: q3a_test2, points: 5, grade: true, solution: false -->
```python
v2 = safe_access({"one": [1, 2, 3], "two": [4, 5, 6], "three": "something"}, "two", [10, 11, 12])
assert set(v2) == set((4, 5, 6))
```

<!-- grade_id: q3a_test3, points: 5, grade: true, solution: false -->
```python
default_val = 3
try:
    value = safe_access({"one": 1, "two": 2}, "three", default_val)
except:
    assert False
else:
    assert value == default_val
```

### Part B

Write a function that finds all occurrences of a specific value in a list, and records--in a separate list--these occurrences in terms of the *indices* where they were found.

Your function should:

 - be named `find_all`
 - take 2 arguments: a list of items, and a single element to search for in the list
 - returns 1 list: a list of *indices* into the input list that correspond to elements in the input list that match what we were looking for
 
For example, `find_all([1, 2, 3, 4, 5, 2], 2)` would return `[1, 5]`. `find_all([1, 2, 3], 0)` should return `[]`.

You cannot use any built-in functions besides `len()` and `range()`.

<!-- grade_id: q3b, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q3b_test1, points: 5, grade: true, solution: false -->
```python
l1 = [1, 2, 3, 4, 5, 2]
s1 = 2
a1 = [1, 5]
assert set(a1) == set(find_all(l1, s1))
```

<!-- grade_id: q3b_test2, points: 5, grade: true, solution: false -->
```python
l2 = ["a", "random", "set", "of", "strings", "for", "an", "interesting", "strings", "problem"]
s2 = "strings"
a2 = [4, 8]
assert set(a2) == set(find_all(l2, s2))
```

<!-- grade_id: q3b_test3, points: 5, grade: true, solution: false -->
```python
l3 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
s3 = 1001
a3 = []
assert set(a3) == set(find_all(l3, s3))
```

### Part C

Write a function which provides counts of very specific elements in a list. This function will return a dictionary where the keys are the elements you wanted to find, and their values are the number of times (counts) they were found in the target list.

*HINT*: You can use your answer to Part B here to expedite the solution! You don't have to, but you can, as it partially solves the problem here. To do that, simply call your function `find_all`--you do NOT need to copy/paste it into the cell below! Python will see it in an earlier cell and will know that it exists (as long as you've hit the "Play/Run" button on that cell, of course. And if you make any changes to it, you should hit the "Play/Run" button again; this updates it in the Python program, kind of like hitting "Save" on a Word doc).

Your function should:

 - be named `element_counts`
 - take 2 arguments, both lists: a list of your data, and a list of elements you want counted in your data
 - return a dictionary: keys are the elements you wanted counted, and values are their counts in the data
 
For example, `element_counts([1, 2, 3, 4, 5, 2], [2, 5])` would return `{2: 2, 5: 1}`, as there were two 2s in the data list, and one 5. Another example would be `element_counts([1, 2, 3], [0, 1])`, which would return `{0: 0, 1: 1}`, as there were no 0s in the original list, and one 1.

You cannot use any built-in functions besides `len()` and `range()`.

<!-- grade_id: q3c, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q3c_test1, points: 5, grade: true, solution: false -->
```python
l1 = [1, 2, 3, 4, 5, 2]
s1 = [2, 5]
a1 = {2: 2, 5: 1}
assert a1 == element_counts(l1, s1)
```

<!-- grade_id: q3c_test2, points: 5, grade: true, solution: false -->
```python
l2 = ["a", "random", "set", "of", "strings", "for", "an", "interesting", "strings", "problem"]
s2 = ["strings", "of", "notinthelist"]
a2 = {"strings": 2, "of": 1, "notinthelist": 0}
assert a2 == element_counts(l2, s2)
```