In this question, you'll be doing some error handling using `try/except` statements.

### Part A

In this question, you'll fill out the code to perform safe division. This means that, in the case of a division-by-zero, you catch the error before it crashes the program, and instead return 0.

`a` and `b` are variables that store the divisor and dividend, respectively. **You CANNOT use conditionals**; instead, you need to use a `try/except` block to gracefully catch any errors resulting from a division by zero.

Assign the final value to the variable `quotient`.

**Hint**: the error you'll want to catch is `ZeroDivisionError`.

**No imports allowed!**

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def safe_divide(a, b):
    quotient = 0

    ### BEGIN SOLUTION
    
    try:
        
        quotient = a / b
        return quotient
        
    except ZeroDivisionError:
        
        quotient = 0
    
    ### END SOLUTION
    
    return quotient
```

<!-- grade_id: q3a_test1, points: 5, grade: true, solution: false -->
```python
a1 = 10
b1 = 5
q1 = 2.0
assert q1 == safe_divide(a1, b1)
```

<!-- grade_id: q3a_test2, points: 10, grade: true, solution: false -->
```python
a2 = 10
b2 = 0
q2 = 0.0
assert q2 == safe_divide(a2, b2)
```

### Part B

In this question, you need to write code that calls the `magic()` function, and returns whatever value `magic` hands back. **Be warned!** This strange `magic` function is prone to crashing and throwing errors, so you'll need to use sufficient `try/except` wizardry to make sure it doesn't bring your program down, too!

In the code below, add a `try/except` block where you call `magic()` and assign it to the variable `retval`. If it doesn't throw an error, just return `retval`. If it *does* throw an error, catch it gracefully, and then assign `-1` to `retval`.

**No imports allowed!**

<!-- grade_id: q3b, points: 0, grade: false, solution: true -->
```python
def do_you_believe_in_magic():
    retval = -1
    
    ### BEGIN SOLUTION
    
    try:
        
        retval = magic()
        
    except:
        
        retval = -1
    
    ### END SOLUTION
    
    return retval
```

<!-- grade_id: q3b_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(91375)

def magic():
    i = np.random.randint(100)
    if i > 50:
        return i >> 2
    else:
        raise ValueError("Something bad happened!")
        
i1 = do_you_believe_in_magic()
assert i1 == 14
```

Output:
```
14
```

<!-- grade_id: q3b_test2, points: 10, grade: true, solution: false -->
```python
np.random.seed(3851)

i2 = do_you_believe_in_magic()
assert i2 == -1

### BEGIN HIDDEN TESTS
i3 = do_you_believe_in_magic()
assert i3 == 24
### END HIDDEN TESTS
```

Output:
```
24
```

### Part C

```python
def modified_sum(a, b, c):
    
    total = 0

    if a == b:
        total = c
        
    elif b == c:
        total = a
        
    elif a == c:
        total = b
    
    else:
        total = a + b + c
        
    return total
        
    

test = modified_sum(1, 2, 3)
print(test)

test = modified_sum(2, 2, 2)
print(test)
```

Output:
```
6
2
```

### Part D

```python
def find13(nums):
    
    length = len(nums) - 1
    thirteen = False
    
    for i in range(length):
        if nums[i] == 1 and nums[i + 1] == 3:
            thirteen = True
    return thirteen

lst = [1,3,2,4]
print(find13(lst))


lst = [1,2,3,4]
print(find13(lst))


lst = [1,2,3,1,3]
print(find13(lst))
```

Output:
```
True
False
True
```

```python

```